# Домашнее задание №1: "Валидация рекомендательных систем"

Правила заполнения ноутбуков на авто-проверку:
- повторить окружение преподавателя. Проверка будет проходить в окружении с python 3.10
```bash
pip install pandas==2.2.3 numpy==1.26.4
```
- писать код только в ячейках с пометкой # YOUR CODE HERE, сразу после этой пометки
- не менять код преподавателя
- не добавлять новые импорты
- не добавлять новые ячейки
- следить, чтобы не было warning - они автоматом фейлят задание
- перед сдачей проверить, что весь ноутбук прогонятся от начала до конца и все тесты проходят

**Важно: в заданиях есть скрытые тесты**

```python
pip install pandas==2.2.3 numpy==1.26.4
```

```python
!python -V
```

```python
# Убедитесь, что вы не добавляете новые импорты в ноутбук. Решение должно быть ограничено данными библиотеками

import warnings
warnings.simplefilter("ignore")

import pandas as pd
import numpy as np
import time
print(pd.__version__)
print(np.__version__)
```

## Кейс 1. Олд-скульный тест на математику (8 баллов)

![01.png](attachment:88596bb5-5036-4539-9c5b-29ab5da3334b.png)

Перед тобой топ 5 рекомендаций от обученного алгоритма для 3-х тестовых юзеров в раскраске позитивам / негативам. У каждого тестового юзера было ровно 2 ground truth интеракции в тестовом периоде. Оранжевые кружки - это айтемы, с которыми юзер взаимодействовал в тестовом периоде, белые - с которыми не взаимодействовал.

Твоя задача - посчитать метрики для оценки качества алгоритма. Сторонние библиотеки для расчёта метрик или код с семинара использовать не стоит - в ранжирующих метриках бывают различия в логике расчётов (см. notes в шпаргалке). Ориентируйся только на формулы в ноутбуке, чтобы ответ сошёлся.
В этом задании лучше не писать универсальные функции для получения ответа, а посчитать на конкретном примере вручную.

![05.png](attachment:328772c4-bc09-4eb5-b7d3-43d55f1f3ca9.png)

Твоя шпаргалка:

-----

${Precision@k} =  \frac{1}{|U|}\sum_{u \in U}\frac{1}{k}\sum_{i=1}^{k} rel_u(i)$

-----

${Recall@k} =  \frac{1}{|U|}\sum_{u \in U}\frac{1}{|R_u|}\sum_{i=1}^{k} rel_u(i)$

-----

$MAP@k = \frac{1}{|U|}\sum_{u \in U}AP_u@k$

$AP_u@k = \frac{1}{k}(\sum_{i=1}^{k} Precision@i * rel_u(i))$

-----

$NDCG@k=\frac{1}{|U|}\sum_{u \in U}\frac{DCG_u@k}{IDCG_u@k}$

$DCG_u@k = \sum_{i=1}^{k} \frac{rel_u(i)}{log_2(i + 1)}$

$IDCG_u@k = \sum_{i=1}^{\min (|R_u|, k)} \frac{1}{log_2(i + 1)}$

-----

- $rel_u(i)$ is an indicator function, it equals to $1$ if the item at rank $i$ is relevant to user $u$, $0$ otherwise
- $|R_u|$ is number of relevant (ground truth) items for user $u$

-----

Note: $AP_u@k$ is not always divided by $k$, sometimes $|R(u)|$ or $min(|R_u|,k)$ is used instead

Note: $IDCG_u@k$ is not always summed for $min(|R_u|,k)$, sometimes $k$ is used instead

Precision учитывает количество релевантных позиций в топе, не учитывая оставшиеся вне топа. Для каждого юзера 1 релевантный айтем.

<!-- grade_id: cell-5e51a0b4b73b5d8d, points: 0, grade: false, solution: true -->
```python
# Классификационные метрики: каждый правильный ответ даёт 1 балл

def precision_at_1() -> float:
    users = [
        [1, 0, 0, 0, 0],
        [0, 1, 0, 0, 0],
        [0, 0, 1, 0, 0]
    ]
    U = len(users)
    pres = sum(map(lambda user: user[0], users))
    return pres/U
    raise NotImplementedError()
    return ans

def precision_at_5() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans

def recall_at_1() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans
```

<!-- grade_id: cell-c3ab127c527aa7a9, points: 0, grade: false, solution: true -->
```python
# Ранжирующие метрики: каждый правильный ответ даёт 1 балл

# AveragePrecision@3 для юзера `2`
def ap_at_3_for_user_2() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans

# MeanAveragePrecision@3
def map_at_3() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans

# DiscountedCumulativeGain@3 для юзера `2`
def dcg_at_3_for_user_2() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans

# IdealDiscountedCumulativeGain@3 для юзера `2`
def idcg_at_3_for_user_2() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans

# NormalizedDiscountedCumulativeGain@3
def ndcg_at_3() -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans
```

### Проверка решений
При проверке преподавателем объём тестов будет расширен

<!-- grade_id: cell-ceafcc6f7cbb7130, points: 1, grade: true, solution: false -->
```python
# Проверяем классификационные метрики
# Правильный ответ даёт 1 балл

assert isinstance(precision_at_1(), float)

```

<!-- grade_id: cell-3d2879d60ca5e96f, points: 1, grade: true, solution: false -->
```python
# Проверяем классификационные метрики
# Правильный ответ даёт 1 балл

assert isinstance(precision_at_5(), float)

```

<!-- grade_id: cell-e892f73bb4a6993e, points: 1, grade: true, solution: false -->
```python
# Проверяем классификационные метрики
# Правильный ответ даёт 1 балл

assert isinstance(recall_at_1(), float)

```

<!-- grade_id: cell-9de65ab3b9f68e43, points: 1, grade: true, solution: false -->
```python
# Проверяем Average Precision (ap_at_3_for_user_2)
# Правильный ответ даёт 1 балл

assert isinstance(ap_at_3_for_user_2(), float)

```

<!-- grade_id: cell-3cf76f15794cdaec, points: 1, grade: true, solution: false -->
```python
# Проверяем Mean Average Precision (map_at_3)
# Правильный ответ даёт 1 балл

assert isinstance(map_at_3(), float)

```

<!-- grade_id: cell-1c735ff73394e6c9, points: 1, grade: true, solution: false -->
```python
# Проверяем Discounted Cumulative Gain (dcg_at_3_for_user_2)
# Правильный ответ даёт 1 балл

assert isinstance(dcg_at_3_for_user_2(), float)

```

<!-- grade_id: cell-a436b938a124b9a3, points: 1, grade: true, solution: false -->
```python
# Проверяем Ideal Discounted Cumulative Gain (idcg_at_3_for_user_2)
# Правильный ответ даёт 1 балл

assert isinstance(idcg_at_3_for_user_2(), float)

```

<!-- grade_id: cell-1f33ca3ece1fb584, points: 1, grade: true, solution: false -->
```python
# Проверяем Normalized Discounted Cumulative Gain (ndcg_at_3)
# Правильный ответ даёт 1 балл

assert isinstance(ndcg_at_3(), float)

```

## Кейс 2. Имплементация взвешенного Recall (12 баллов)

![02.png](attachment:bd952d1c-1a30-47c2-8cf4-0d54949caeb5.png)

Ты разрабатываешь полку с персональными рекомендациями в интернет-магазине.
Качество рекомендаций важно оценивать в деньгах.
Имплементируй функцию, которая рассчитывает взвешенный Recall на основе рекомендаций модели и реальных покупок пользователей.

Метрика усреднена по юзерам и идентична обычному Recall, но релеватность айтема для юзера оценивает не в бинарном виде (hit / miss), а в денежном. Какую долю реально потраченных пользователями денег, в среднем умеет предсказывать модель?

```python
# Отранжированные рекомендации товаров, полученные из обучения алгоритма на исторических данных

top_k = 5
n_users = 3
reco = pd.DataFrame({
    "user_id": np.repeat(np.arange(n_users), top_k),
    "item_id": np.tile(np.arange(top_k), n_users),
})
reco["rank"] = reco.groupby("user_id").cumcount() + 1
reco
```

```python
# Реальные покупки пользователей в тестовом периоде

test = pd.DataFrame({
    "user_id": [0, 1, 2, 0, 1, 2],
    "item_id": [0, 1, 2, 3, 4, 5],
    "money": [100, 200, 100, 200, 100, 500]
})
test
```

Имплементируйте функцию для расчёта взвешенного Recall. Помните, что метрика сначала считается для каждого юзера, а потом усредняется для получения финального значения.  
Вспомните подходы с семинара, которые позволяют рассчитывать метрики в векторном виде.  
Корректная, но неэффективная имплементация даст 2 балла из 5.  
Корректная и эффективная - 5 из 5.

<!-- grade_id: cell-8d3a9fc17ce25556, points: 0, grade: false, solution: true -->
```python
def weighted_recall(reco: pd.DataFrame, test: pd.DataFrame, k: int, weight_col: str = "money") -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
    return res
```

### Проверка решения
При проверке преподавателем объём тестов будет расширен

<!-- grade_id: cell-c27e3c2f6a5ea3c1, points: 6, grade: true, solution: false -->
```python
# Проверяем корректность расчётов
# Функция с корректными расчётами даёт 6 баллов
# Хардкод получит 0 баллов при проверке преподавателем

actual = weighted_recall(reco.copy(), test.copy(), k=5)
expected = (1 + 1 + 1 / 6) / 3
assert isinstance(actual, float)
assert abs(expected - actual) < 0.001

actual = weighted_recall(reco.copy(), test[test["user_id"]==0].copy(), k=5)
expected = 1
assert isinstance(actual, float)
assert abs(expected - actual) < 0.001

```

<!-- grade_id: cell-dc553e3f7361a4c8, points: 6, grade: true, solution: false -->
```python
# Проверяем эффективность имплементации
# Эффективная имплементация в векторной форме с корректными расчётами даёт 6 баллов
# Хардкод получит 0 баллов при проверке преподавателем
# Наивная имплементация получит 0 баллов при проверке преподавателем

class Timer:
    def __enter__(self):
        self._start = time.perf_counter()
        self._end = None
        return self

    def __exit__(self, *args):
        self._end = time.perf_counter()

    @property
    def elapsed(self):
        return self._end - self._start


with Timer() as timer:
    actual = weighted_recall(reco, test, k=5)
assert isinstance(actual, float)
assert timer.elapsed < 0.1

```

# Как сдать ноутбук на проверку
1. Прогоните весь код ноутбука - проверьте, что нет ошибок, весь код запускается и тесты проходят
2. Обязательно! Для кейса 2 проверьте, что ваш код будет корректно работать на новых данных
3. Приложите готовый ноутбук при сдаче домашнего задания

**Максимальная оценка: 20 баллов**